```python
# %load ./test/test_floatcheck.py
```

```python
# %load ./test/test_vectorcheck.py
```

```python
# %load ./test/test_matrixcheck.py
```

In order to successfully complete this assignment you need to participate both individually and in groups during class on **###DUE_DATE###**.


# In-Class Assignment: Diagonalization

<img src="https://wikimedia.org/api/rest_v1/media/math/render/svg/62ab0ef52ecb1e1452efe6acf096923035c75f62" width="50%">

Image from: [https://en.wikipedia.org/wiki/Diagonalizable_matrix](https://en.wikipedia.org/wiki/Diagonalizable_matrix)

    

### Agenda for today's class (80 minutes)

</p>

###TOC###

1. [(20 minutes) Homework 3 Review](#Homework_3_Review)
1. [(20 minutes) Pre-class Assignment Review](#Pre-class_Assignment_Review)
1. [(20 minutes) Diagonalization](#Diagonalization)
1. [(20 minutes) The Power of a Matrix](#The_Power_of_a_Matrix)


```python

```

----

# 0. Testing new instructional software

Today in class we will be testing out some new instructional software. We hope to use this software as we increase the number of students in the class from 1 section of 50 this semester to 4 sections of 50 (200 students!!) in the fall. The software has two goals:

1. Provide immediate feedback to students as they are learning new skills. We hope this immediate feedback will help students gain confidence in their understanding and provide support for self learning. 
2. Reduce the amount of time and effort for instructors to grade assignments so that they can spend more time working with students. 

This assignment uses the new software.  We ask that you give it a try and give us feedback we can use to improve the student experience.  There are additional instructions at the end of the assignment to turn it in.  

The visiting instructors conducting this experiment have no impact on your grade and anything you tell them will not impact your grade. 

```python
from urllib.request import urlretrieve
urlretrieve("https://raw.githubusercontent.com/colbrydi/jupytercheck/master/answercheck.py", "answercheck_old.py");
```

```python
from answercheck import checkanswer
checkanswer.detailedwarnings = False
```

Here is an example demonstrating how ```checkanswer``` works. The "magic string" in the function (2e55d74f5c78981f6b877b198bdc61ba) should not be changed. This is the "hash" function that secretly encodes the answer.

```python
f = 0.0111
checkanswer.float(f, '2e55d74f5c78981f6b877b198bdc61ba');
```

Output:
```
Testing 0.0111
Answer seems to be correct
```

Notice if you give it the wrong answer it gives you an error

```python
f = 4
checkanswer.float(f, '2e55d74f5c78981f6b877b198bdc61ba');
```

We are trying to make the ```checkanswer``` functions smart and handle correct answers that are in different formats. For example the following is equivalent within a reasonable roundoff and therefore true. 

```python
f = 0
checkanswer.float(f, '30565a8911a6bb487e3745c0ea3c8224');
```

Output:
```
CheckWarning: passed variable is <class 'int'> and not a float. 
    Trying to convert to a float using ```A = float(A)```.


CheckWarning: Value is negative zero...
    Converting to positive zero before checking using ```A = 0.00```.

Testing 0.0
Answer seems to be correct
```

```python
f = -0
checkanswer.float(f, '30565a8911a6bb487e3745c0ea3c8224');
```

Output:
```
CheckWarning: passed variable is <class 'int'> and not a float. 
    Trying to convert to a float using ```A = float(A)```.


CheckWarning: Value is negative zero...
    Converting to positive zero before checking using ```A = 0.00```.

Testing 0.0
Answer seems to be correct
```

```python
f = 1
checkanswer.float(f, 'e4c2e8edac362acab7123654b9e73432');
```

Output:
```
CheckWarning: passed variable is <class 'int'> and not a float. 
    Trying to convert to a float using ```A = float(A)```.

Testing 1.0
Answer seems to be correct
```

```python
f = 0.011100001
checkanswer.float(f, '2e55d74f5c78981f6b877b198bdc61ba');
```

Output:
```
Testing 0.0111
Answer seems to be correct
```

---
<a name="Homework_3_Review"></a>
# 1. Homework 3 Review

* [0228-HW3-Networks_Spectra_and_Jurassic_Park.ipynb](0228-HW3-Networks_Spectra_and_Jurassic_Park.ipynb)

---
<a name="Pre-class_Assignment_Review"></a>
# 2. Pre-class Assignment Review

* [0309--Diagonalization-pre-class-assignment.ipynb](0309--Diagonalization-pre-class-assignment.ipynb)

----
<a name="Diagonalization"></a>
# 3. Diagonalization

## The eigenvalues of triangular (upper and lower) and diagonal matrices

* The eigenvalues for triangular matrices are the diagonal elements.
* The eigenvalues for the digonal matrices are the diagonal elements. 

## Diagonalization


**Definition**: A square matrix $A$ is said to be *diagonalizable* if there exist a matrix $C$ such that $D=C^{-1}AC$ is a diagonal matrix.

**Definition**: $B$ is a *similar matrix* of $A$ if we can find $C$ such that $B=C^{-1}AC$.


Given an $n\times n$ matrix $A$, can we find another $n \times n$ invertable matrix $C$ such that when $D=C^{-1}AC$ is diagonal, i.e., $A$ is diagonalizable?

* Because $C$ is inveritble, we have 
$$C^{-1}AC=D \\ CC^{-1}AC = CD\\ AC = CD $$


* Generate $C$ as the columns of $n$ linearly independent vectors $(x_1...x_n)$ We can compute $AC=CD$ as follows:
$$ A\begin{bmatrix} \vdots  & \vdots  & \vdots  & \vdots  \\ \vdots  & \vdots  & \vdots  & \vdots  \\ { x }_{ 1 } & { x }_{ 2 } & \dots  & { x }_{ n } \\ \vdots  & \vdots  & \vdots  & \vdots  \end{bmatrix}=AC=CD=\begin{bmatrix} \vdots  & \vdots  & \vdots  & \vdots  \\ \vdots  & \vdots  & \vdots  & \vdots  \\ { x }_{ 1 } & { x }_{ 2 } & \dots  & { x }_{ n } \\ \vdots  & \vdots  & \vdots  & \vdots  \end{bmatrix}\begin{bmatrix} { \lambda  }_{ 1 } & 0 & 0 & 0 \\ 0 & { \lambda  }_{ 2 } & 0 & 0 \\ \vdots  & \vdots  & { \dots  } & \vdots  \\ 0 & 0 & 0 & { \lambda  }_{ n } \end{bmatrix}$$
* Then we check the corresponding columns of the both sides. We have 
$$Ax_1 = \lambda_1x_1\\\vdots\\Ax_n=\lambda x_n$$

* $A$ has $n$ linear independent eigenvectors.

* $A$ is saied to be *similar* to the diagonal matrix $D$, and the transformation of $A$ into $D$ is called a *similarity transformation*.

### A simple example

Consider the following:
$$ A = \begin{bmatrix}7& -10\\3& -4\end{bmatrix},\quad C = \begin{bmatrix}2& 5\\1& 3\end{bmatrix}$$

&#9989;  **<font color=red>Do this:</font>** Find the similar matrix $D = C^{-1}AC$ of $A$.

```python
#Put your answer to the above question here.
```

```python
##ANSWER##
A = np.matrix([[7, -10],[3,-4]])
vals, vects = np.linalg.eig(A)
C = vects
D = C**-1*A*C
##ANSWER##
```

<!-- grade_id: cell-3cdb9915439d45fe, points: 5, grade: true, solution: false -->
```python
checkanswer.matrix(D, '8313fe0f529090d6a8cdb36248cfdd6c');
```

&#9989;  **<font color=red>Do this:</font>** Find the eigenvalues and eigenvectors of $A$. Set variables ```e1``` and ```vec1``` to be the smallest eigenvalue and it's associated eigenvector and ```e2, vec2``` to represent the  largest.

```python
#Put your answer to the above question here.
```

```python
# ##ANSWER##
# ##Testing using sympy
# eig = A.eigenvals(multiple=True)
# vec = A.eigenvects(multiple=False)


# e1 = vec[0][0]
# e2 = vec[1][0]

# vec1 = vec[0][2]
# vec2 = vec[1][2]

# vec1_1 = vec1
# vec2_2= vec2
# ##ANSWER##
```

```python
vals
```

```python
e1 = vals[1]
e2 = vals[0]
```

<!-- grade_id: cell-f4fda102502f50f9, points: 5, grade: true, solution: false -->
```python
checkanswer.float(e1, "e4c2e8edac362acab7123654b9e73432");
```

<!-- grade_id: cell-88300f29b8aec498, points: 5, grade: true, solution: false -->
```python
checkanswer.float(e2, "d1bd83a33f1a841ab7fda32449746cc4");
```

```python
vec1 = vects[:,1]
vec2 = vects[:,0]
```

<!-- grade_id: cell-f26e2f5a3e41bdd8, points: 5, grade: true, solution: false -->
```python
checkanswer.vector(vec1, "9811eb81debc5939e28ad8c2455169c5")
```

<!-- grade_id: cell-a0ef501c592a3fcc, points: 5, grade: true, solution: false -->
```python
checkanswer.vector(vec2, "d76afb3b2331d589fe5a4172d98808be")
```

```python
##ANSWER##
#Same solution using numpy
A = np.matrix([[7.0, -10],[3,-4]])
C = np.matrix([[2.0, 5],[1,3]])
D = C**-1*A*C
sym.Matrix(D)

vals,vecs = np.linalg.eig(np.matrix(A))

e2 = vals[0]
e1 = vals[1]

vec2 = vecs[:,0]
vec1 = vecs[:,1]
##ANSWER##
```

```python
import numpy as np
A = np.matrix([[1.0, 0.0], [0.0, 2.0]])
checkanswer.matrix(A, '943d90b283d21136da008160e002a5ce');
```

```python
##ANSWER##
checkanswer.float(e1, "e4c2e8edac362acab7123654b9e73432");
##ANSWER##
```

```python
##ANSWER##
checkanswer.float(e2, "d1bd83a33f1a841ab7fda32449746cc4");
##ANSWER##
```

```python
A = [1.0, 0.0, 2.0]
checkanswer.vector(A, '4c50a36b5ea58ab121d70de1b47134a5');
    
```

```python
A = np.matrix([[1.0, 0.0, 2.44]])
checkanswer.vector(A, 'adc7e01b8bd1778081f6ee348b5eac8c');
```

```python
##ANSWER##
checkanswer.vector(vec1, "9811eb81debc5939e28ad8c2455169c5")
##ANSWER##
```

```python
##ANSWER##
checkanswer.vector(vec2, "d76afb3b2331d589fe5a4172d98808be")
##ANSWER##
```

**Theorem:** Similar matrices have the same eigenvalues.

**Proof:** Assume $B=C^{-1}AC$ is a similar matrix of $A$, and $\lambda$ is an eigenvalue of $A$ with corresponding eigenvector $x$. That is, $$Ax=\lambda x$$ 
Then we have $$B(C^{-1}x) = C^{-1}AC(C^{-1}x) = C^{-1}Ax = C^{-1}(\lambda x)= \lambda (C^{-1}x).$$
That is $C^{-1}x$ is an eigenvector of $B$ with eigenvalue $\lambda$.     

### A second example

&#9989;  **<font color=red>Do this:</font>** Consider 
$$ A = \begin{bmatrix}-4& -6\\3& 5\end{bmatrix}.$$
Find a matrix $C$ such that $C^{-1}AC$ is diagonal. (Hint, use the function `diagonalize` in `sympy`.)

```python
#Put your answer to the above question here. 
```

```python
##ANSWER##
A = sym.Matrix([[-4,-6],[3,5]])
C, D = A.diagonalize()
print(C)
print(D)
##ANSWER##
```

<!-- grade_id: cell-d9c7ff4aa895199e, points: 5, grade: true, solution: false -->
```python
#Check the output type
assert(type(C)==sym.Matrix)
```

<!-- grade_id: cell-2c06b41f80b7a258, points: 5, grade: true, solution: false -->
```python
checkanswer.matrix(C,'ba963b7fef354b4a7ddd880ca4bac071')
```

```python
##ANSWER##
A = sym.Matrix([[-4,-6],[3,5]])
A.eigenvals()
A.eigenvects()
C = A.eigenvects()[0][2][0]
C = C.col_insert(1, A.eigenvects()[1][2][0])
C = sym.Matrix(C)
C.inv()*A*C
print(C)
##ANSWER##
```

<!-- grade_id: cell-6a1b21e5f63f4497, points: 5, grade: true, solution: false -->
```python
##ANSWER##
checkanswer.matrix(C,'ba963b7fef354b4a7ddd880ca4bac071')
##ANSWER##
```

### The third example

&#9989;  **<font color=red>Do this:</font>** Consider 
$$ A = \begin{bmatrix}5& -3\\3& -1\end{bmatrix}.$$
Can we find a matrix $C$ such that $C^{-1}AC$ is diagonal?  (Hint: find eigenvalues and eigenvectors using `sympy`)

<!-- grade_id: cell-8eb9fd1f4a5a6136, points: 0, grade: true, solution: true -->
```python
#Put your answer to the above question here. 
```

```python
##ANSWER##
A = sym.Matrix([[5,-3],[3,-1]])
C, D = A.diagonalize()
print(C)
print(D)

## Matrix is not diagonalizable
##ANSWER##
```

## Dimensions of eigenspaces and diagonalization

**Definition**: The set of all eigenvectors of a $n\times n$ matrix corresponding to a eigenvalue $\lambda$, together with the zero vector, is a subspace of $R^n$. This subspace spaces is called *eigenspace*.

* For the third example, we have that the characteristic equation $(\lambda-2)^2=0$.
* Eigenvalue $\lambda=2$ has multiplicity 2, but the eigenspace has dimension 1, since we can not find two lineare independent eigenvector for $\lambda =2$. 

> The dimension of an eigenspace of a matrix is less than or equal to the multiplicity of the corresponding eigenvalue as a root of the characteristic equation.

> A matrix is diagonalizable if and only if the dimension of every eigenspace is equal to the multiplicity of the corresponding eigenvalue as a root of the characteristic equation.

### The fourth example

&#9989;  **<font color=red>Do this:</font>** Consider 
$$ A = \begin{bmatrix}2& -1\\1& 2\end{bmatrix}.$$
Can we find a matrix $C$ such that $C^{-1}AC$ is diagonal?

<!-- grade_id: cell-3bc59d8f51537cae, points: 0, grade: true, solution: true -->
```python
#Put your answer to the above question here. 
```

```python
##ANSWER##
A = sym.Matrix([[2,-1],[1,2]])
C, D = A.diagonalize()

print(C)
print(D)
##ANSWER##
```

---

<a name="The_Power_of_a_Matrix"></a>
# 4. The Power of a Matrix

* For a diagonalizable matrix $A$, we have $C^{-1}AC=D$. Then we have 
$$A = C D C^{-1}$$
* We have 
$$A^2 = C D C^{-1} C D C^{-1} = C D^2 C^{-1}$$
$$A^n = C D C^{-1} \dots C D C^{-1} = C D^n C^{-1}$$
* Because the columns of $C$ are eigenvectors, so we can say that the eigenvectors for $A$ and $A^n$ are the same if $A$ is diagonalizable. 
* If $x$ is an eigenvector of $A$ with the corresponding eigenvalue $\lambda$, then $x$ is also an eigenvector of $A^n$ with the corresponding eigenvalue $\lambda^n$.

```python
# Here are some libraries you may need to use
%matplotlib inline
import numpy as np
import sympy as sym
import networkx as nx
import matplotlib.pyplot as plt
sym.init_printing(use_unicode=True)
```

### Graph Random Walk

* Define the following matrices:
    * $I$ is the identity matrix
    * $A$ is the adjacency matrix
    * $D$ is diagonal matrix of degrees (number of edges connected to each node)
    
$$W=\frac{1}{2}(I + AD^{-1})$$

* The **lazy random walk matrix**, $W$, takes a distribution vector of *stuff*, $p_{t}$, and diffuses it to its neighbors:

$$p_{t+1}=Wp_{t}$$

* For some initial distribution of *stuff*, $p_{0}$, we can compute how much of it would be at each node at time, $t$, by powering $W$ as follows:

$$p_{t}=W^{t}p_{0}$$

* Plugging in the above expression yields:

$$p_{t}=\frac{1}{2}(I + AD^{-1})^{t}p_{0}$$

<font color=red>**DO THIS**</font>: Using matrix algebra, show that $\frac{1}{2}(I + AD^{-1})$ is **similar** to  $I-\frac{1}{2}N$, where $N=D^{-\frac{1}{2}}(D-A)D^{-\frac{1}{2}}$ is the normalized graph Laplacian. 

<!-- grade_id: cell-1a93e034adef3eb1, points: 0, grade: true, solution: true -->
**Your answer goes here** (follow along after attempting)

### Random Walk on Barbell Graph

To generate the barbell graph, run the following cell.

```python
n = 60 # number of nodes
B = nx.Graph() # initialize graph

## initialize empty edge lists
edge_list_complete_1 = [] 
edge_list_complete_2 = []
edge_list_path = []

## generate node lists
node_list_complete_1 = np.arange(int(n/3))
node_list_complete_2 = np.arange(int(2*n/3),n)
node_list_path = np.arange(int(n/3)-1,int(2*n/3))

## generate edge sets for barbell graph
for u in node_list_complete_1:
    for v in np.arange(u+1,int(n/3)):
        edge_list_complete_1.append((u,v))
        
for u in node_list_complete_2:
    for v in np.arange(u+1,n):
        edge_list_complete_2.append((u,v))

for u in node_list_path:
    edge_list_path.append((u,u+1))

# G.remove_edges_from([(3,0),(5,7),(0,7),(3,5)])

## add edges
B.add_edges_from(edge_list_complete_1)
B.add_edges_from(edge_list_complete_2)
B.add_edges_from(edge_list_path)


## draw graph
pos=nx.spring_layout(B) # positions for all nodes

### nodes
nx.draw_networkx_nodes(B,pos,
                       nodelist=list(node_list_complete_1),
                       node_color='c',
                       node_size=400,
                       alpha=0.8)
nx.draw_networkx_nodes(B,pos,
                       nodelist=list(node_list_path),
                       node_color='g',
                       node_size=200,
                       alpha=0.8)
nx.draw_networkx_nodes(B,pos,
                       nodelist=list(node_list_complete_2),
                       node_color='b',
                       node_size=400,
                       alpha=0.8)


### edges
nx.draw_networkx_edges(B,pos,
                       edgelist=edge_list_complete_1,
                       width=2,alpha=0.5,edge_color='c')
nx.draw_networkx_edges(B,pos,
                       edgelist=edge_list_path,
                       width=3,alpha=0.5,edge_color='g')
nx.draw_networkx_edges(B,pos,
                       edgelist=edge_list_complete_2,
                       width=2,alpha=0.5,edge_color='b')


plt.axis('off')
plt.show() # display
```

&#9989;  **<font color=red>Do this</font>:** Generate the lazy random walk matrix, $W$, for the above graph.

```python
#Put your answer to the above question here.
A = nx.adjacency_matrix(B)
A = A.todense()

d = np.sum(A,0) # Make a vector of the sums.
D = np.diag(np.asarray(d)[0])

```

```python
##ANSWER##
A = nx.adjacency_matrix(B)
A = A.todense()

d = np.sum(A,0) # Make a vector of the sums.
D = np.diag(np.asarray(d)[0])

W = 0.5*(np.matrix(np.eye(n)) + np.matrix(A)*np.matrix(D)**(-1))
##ANSWER##
```

<!-- grade_id: cell-fb79da016761443e, points: 5, grade: true, solution: false -->
```python
checkanswer.matrix(W, "8ba624169e3a2ab216081049e1429eae")
```

&#9989;  **<font color=red>Do this</font>:** Compute the eigenvalues and eigenvectors of $W$. Make a diagonal matrix $A$ with the eigenvalues on the diagonal. Name the matrix of eigenvectors $V$ (each column is an eigenvector). (**Note**: We now have that $W=VAV^{-1}=VAV^{T}$ since $V$ is unitary.)

```python
#Put your answer to the above question here. 
```

```python
##ANSWER##
lam, vecs =  np.linalg.eig(W)

A = np.diag(lam)
V = vecs


np.allclose(V*A*V**(-1),W)
##ANWER##
```

<!-- grade_id: cell-10f3b45c444ad7e4, points: 5, grade: true, solution: false -->
```python
checkanswer.matrix(A, 'e82af1c40ba1e7bb14594e41aa3d0752')
```

```python
V[V==-0]
```

<!-- grade_id: cell-2f59a9e06170dac5, points: 5, grade: true, solution: false -->
```python
checkanswer.matrix(V, '9ae51df03acd5a0fa71ee791c0964812')
```

&#9989;  **<font color=red>Do this</font>:** Let your $p_{0}=[1,0,0,\ldots,0]$. Compute $p_{t}$ for $t=1,2,\ldots,100$, and plot $||v_{1}-p_{t}||_{1}$ versus $t$, where $v_{1}$ is the eigenvector associated with the largest eigenvalue $\lambda_{1}=1$ and whose sum equals 1. (**Note**: $||\cdot||_{1}$ may be computed using ```np.linalg.norm(v_1-p_t, 1)```.)

<!-- grade_id: cell-9e691ac811c35e4d, points: 5, grade: true, solution: true -->
```python
#Put your answer to the above question here. 
```

```python
##ANSWER##
p0 = np.zeros([60,1])
p0[0] = 1
p0 = np.matrix(p0)

v1 = V[:,0]
v1 = v1/sum(v1)  #  normalize the eigenvector such that is sums one

err = np.zeros(10000)
p = p0
for i in range(0,10000):
    p = W*p
    err[i] = np.linalg.norm(v1-p, 1)

plt.semilogy(err)
##ANSWER##
```

#### Compare to Complete Graph

If you complete the above, do the same for a complete graph on the same number of nodes.

&#9989;  **<font color=red>Question</font>:** What do you notice about the graph that is different from that above?

<!-- grade_id: cell-9cadbdd3014757bc, points: 5, grade: true, solution: true -->
Put your answer to the above question here.

---------
### Congratulations, we're done!

Now, you just need to submit this assignment by uploading it to the course <a href="https://d2l.msu.edu/">Desire2Learn</a> web page for today's dropbox (Don't forget to add your names in the first cell).


### Course Resources:

- [Syllabus](###SYLLABUS###)
- [Preliminary Schedule](###SCHEDULE###)
- [Course D2L Page](###D2L###)

```python
##ANSWER## 
#this cell gets the name of the current notebook.
from jupyterinstruct import makestudent 
makestudent.getname()

import thiscourse
tags = thiscourse.tags()
tags['YEAR'] = '2020'
##ANSWER## 
```

```python
##ANSWER## 
#This cell runs the converter which removes ANSWER feilds, renames the notebook and cleans out output fields. 
makestudent.merge(this_notebook, "./"+tags['COURSE_CODE']+"/", tags)
##ANSWER## 
```

```python
##ANSWER##
#This cell runs the nbgrader command which copies the files to a module sorce directory and then generates a release code
makestudent.usenbgrader(this_notebook, f"./{tags['COURSE_CODE']}/", tags)
##ANSWER##
```

&#169; Copyright ###YEAR###,  Michigan State University Board of Trustees